# Denison CS181/DA210 SW Lab #2 - Step 1

Before you get your checkpoint, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells** (in the menubar, select Kernel$\rightarrow$Restart And Run All).

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

```python
import os
import os.path
import io
import sys

datadir = "publicdata"
```

---
## Part A: Dictionaries

A dictionary allows us to represent key-value pairs in Python in a way that ensures fast lookup of values.

In Python, dictionaries are represented using "curly braces": `{}`.  (This is in contrast to `[]` used by lists or `()` used by tuples.)

An empty dictionary is thus given by:

```python
emptyDict = {}
```

We can statically create a dictionary by specifying the key-value pairs.  Keys can be any immutable types (e.g., strings, ints, tuples).  Values can be *any* types (e.g., strings, ints, tuples, lists, floats, other dictionaries).

For example, we can map people's names to their favorite colors:

```python
favoriteColors = {"Isabella": "teal",
                  "Miyeon": "purple",
                  "Joshua": "brown",
                  "Abhishek": "blue"}
```

We can use square brackets, `[]`, to look up the value corresponding to a key:

```python
# Look up the value for a given key
favoriteColors["Isabella"]
```

We can also use square brackets, `[]`, to add a new key-value pair to the dictionary.

```python
print(favoriteColors)

# Add a new key-value pair
favoriteColors["Enrico"] = "yellow"

print(favoriteColors)
```

If we provide an existing key, we can update its value.  Note that this means that _there cannot be duplicate keys_ in a Python dictionary.

```python
print(favoriteColors)

# Update a key-value pair
favoriteColors["Enrico"] = "orange"

print(favoriteColors)
```

We can use the `in` keyword to check if a key exists in the dictionary:

```python
print("Isabella" in favoriteColors)

print("Lulu" in favoriteColors)
```

Similarly, we can use a `for` loop to iterate through the keys in a dictionary:

```python
for key in favoriteColors:
    value = favoriteColors[key]
    print(key + "'s favorite color is " + value)
```

We can save a line in this for loop similar to using `enumerate()` with a list's indices and elements:

```python
for key, value in favoriteColors.items():
    print(key + "'s favorite color is " + value)
```

> You've reached the first checkpoint in the lab.  Make sure to have it signed off by the instructor or TA.
>
> Checkpoint 1: Does `D = {'a': "ant", 'c': "cat", 'a': "alligator"}` represent a legitimate dictionary?  Explain.

---

## Part B

**Q1** Create a dictionary to represent the following data.  The keys should be strings and the values should be integers.

Name       | Age     
:--------- |:-------------
Lulu       | 9
Pug        | 15
Max        | 7
Sadie      | 14
Hobbes     | 9

You should name your dictionary `petAgeMap`.

```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Testing cell
assert len(petAgeMap) == 5

assert "Sadie" in petAgeMap
assert "Hobbes" in petAgeMap
assert "Lulu" in petAgeMap
assert petAgeMap["Sadie"] == 14
assert petAgeMap["Hobbes"] == 9
assert petAgeMap["Lulu"] == 9
```

**Q2** Write a function

```
    incrementValue(d, key)
```

that takes a dictionary and key and increments the corresponding value by one.  Your function should modify the dictionary in place and not return anything.  You may assume that dictionary values are integers and that the key already exists in the dictionary.

```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Testing cell
assert incrementValue.__doc__ is not None # don't forget the docstring!

assert len(petAgeMap) == 5
assert petAgeMap["Lulu"] == 9
incrementValue(petAgeMap, "Lulu")
assert len(petAgeMap) == 5
```

---
## Part C

**Q3** Write a function

```
    makeDictFromList(L)
```
    
that creates a dictionary whose keys are the items of `L` and whose values are the corresponding *indices* of `L`.  You may assume `L` is a list without any duplicate items.

<!-- grade_id: cell-1b48e481c8c4c6cc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7a90d00c4694cb19, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert makeDictFromList.__doc__ is not None # don't forget the docstring!

L1 = [5, 4, 3, 2, 1]
expect1 = {5: 0, 4: 1, 3: 2, 2: 3, 1:4}
assert makeDictFromList(L1) == expect1

L2 = "Piotr Piper Picked a Pack of Pickled Peppers".split()
expect2 = {'Piotr': 0, 'Piper': 1, 'Picked': 2, "a": 3, "Pack": 4, "of": 5, "Pickled": 6, "Peppers": 7}
assert makeDictFromList(L2) == expect2
```

> You've reached the second checkpoint in the lab.  Make sure to have it signed off by the instructor or TA.
>
> Checkpoint 2: Given your implementation of `makeDictFromList`, what would happen if there were duplicate entries?  How could you instead represent multiple values for one key?

**Q4** Write a function

```
    filmography(L)
```

whose parameter `L` is a list of tuples of the form `(actor, movie)`, and creates and returns a dictionary where the keys are actor names and the values are lists of movies that actor has been in.  For example, given 

```
    L = [(‘DiCaprio’,‘Revenant’),
         (‘Cumberbatch’,‘Hobbit’),
         (‘DiCaprio’,‘Gatsby’),
         (‘DiCaprio’,‘Django’)]
```

then you should return 

```
    {‘DiCaprio’:[‘Revenant’,‘Gatsby’,‘Django’],
     ‘Cumberbatch’:[‘Hobbit’]}
```

<!-- grade_id: cell-6b2af6f816026d9f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Sanity check
L = [('DiCaprio','Revenant'),
     ('Cumberbatch','Hobbit'),
     ('DiCaprio','Gatsby'),
     ('DiCaprio','Django')]
D = filmography(L)
print(D)
```

<!-- grade_id: cell-687b64ee889a4671, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert filmography.__doc__ is not None # don't forget the docstring!

L = [('DiCaprio','Revenant'),
     ('Cumberbatch','Hobbit'),
     ('DiCaprio','Gatsby'),
     ('DiCaprio','Django')]
D = filmography(L)
assert len(D) == 2
assert 'DiCaprio' in D.keys()
assert 'Django' in D['DiCaprio']
assert 'Cumberbatch' in D.keys()
assert D['Cumberbatch'] == ['Hobbit']
```