# Denison CS181/DA210 SW Lab #2 - Step 2

Before you get your checkpoint, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells** (in the menubar, select Kernel$\rightarrow$Restart And Run All).

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import os
import os.path

datadir = "publicdata"
```

---

## Part D

**Q1** As a warm-up, fill in the body of the `for` loop below to accumulate data from the rows in `"topnames.csv"` into the parallel lists initialized for you as `yearList`, `sexList`, etc.

Here is a snippet of `topnames.csv`:

```
year,sex,name,count
...
2016,Male,Noah,19117
2016,Female,Emma,19496
2017,Female,Emma,19800
2017,Male,Liam,18798
2018,Female,Emma,18688
2018,Male,Liam,19837
```

```python
path = os.path.join(datadir, "topnames.csv")

yearList = []
sexList = []
nameList =[]
countList = []

with open(path, "r") as f:
    headers = f.readline().strip().split(',') # get the headers in a list
    print(headers)                            # see what the headers are
    for line in f:
        # YOUR CODE HERE
        raise NotImplementedError()
```

```python
# Sanity checks
print(yearList[:6])
print(sexList[:6])
print(nameList[:6])
print(countList[:6])

print()
print(yearList[-1], sexList[-1], nameList[-1], countList[-1])
```

---

## Part E: Dictionary of Lists (DoL)

In this representation, the keys are the _names of the columns_.  The values being mapped to are _lists_.

Each entry in the DoL represents **exactly one column** of the data set.  For example, in `topnames.csv`, column names are `"year"`, `"sex"`, `"name"`, and `"count"`.

Elements in the column list represent the column value for the successive rows of the data set.

Here is an example hard-coded DoL corresponding to the `topnames.csv` snippet above:

```python
topnames = {'year':  [2018, 2018, 2017, 2017, 2016, 2016],
            'sex':   ['Male', 'Female', 'Male', 'Female', 'Male', 'Female'],
            'name':  ['Liam', 'Emma', 'Liam', 'Emma', 'Noah', 'Emma'],
            'count': [19837, 18688, 18798, 19800, 19117, 19496]}
```

When a representation makes it easy to work with entire columns, but working with rows is more cumbersome, it is a **column-centric** representation.

Here is how you can get a column of data with a DoL:

```python
# Column operation: all values of 'count' column
print(topnames["count"])
```

It is much more work to get a row of data using a DoL:

```python
# Row operation: all column values for row with index 3
row_index = 3
print(topnames['year'][row_index], topnames['sex'][row_index],
      topnames['name'][row_index], topnames['count'][row_index])
```

It is very important that the values in each inner list (each column) correspond properly to each other.  If we delete an element of one list, the data won't line up:

```python
# Breaking correspondence
del topnames["sex"][0]

# Now Emma is a "male" name
row_index = 3
print(topnames['year'][row_index], topnames['sex'][row_index],
      topnames['name'][row_index], topnames['count'][row_index])
```

> You've reached the third checkpoint in the lab.  Make sure to have it signed off by the instructor or TA.
>
> Checkpoint 3: What expression could you use to get the number of rows in `topnames` (assuming we hadn't modified it in the previous cell).  What about the number of columns?

---

## Part F

**Q2** Assuming that `path` refers to a CSV file with the header line of `year,sex,name,count` and data lines with those same four fields, write a function

```
    readTopNamesDoL(path)
```

that reads the file to create a DoL representation and returns that dictionary from the function.

<!-- grade_id: cell-8849b03d485bbf5d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Printing cell
tn10 = readTopNamesDoL(os.path.join(datadir, "tn10.csv"))
print(tn10)
```

<!-- grade_id: cell-44470aa818a11e56, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert readTopNamesDoL.__doc__ is not None # don't forget the docstring!

tn10 = readTopNamesDoL(os.path.join(datadir, "tn10.csv"))
assert isinstance(tn10, dict)
assert len(tn10) == 4
assert 'year' in tn10
assert 'sex' in tn10
assert 'count' in tn10
assert len(tn10['year']) == 10
```

**Q3** Write a function 

```
    filterTopNamesDoL(tnDoL, threshold)
```
    
to create a **filter** copy of a topnames DoL `tnDoL` (with columns `year`, `sex`, `name`, `count`) so that only rows with a `count` value greater than or equal to `threshold` are present in the newly created DoL.

Your function should return the new and filtered DoL, and not modify the original `tnDoL`.

<!-- grade_id: cell-698093a1af18a3b9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Printing cell
tn10 = readTopNamesDoL(os.path.join(datadir, "tn10.csv"))
tn10_filter = filterTopNamesDoL(tn10, 19000)
print(tn10_filter)
```

<!-- grade_id: cell-a61460d2e6253eef, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert filterTopNamesDoL.__doc__ is not None # don't forget the docstring!

topnames = {'year':  [2018, 2018, 2017, 2017, 2016, 2016],
            'sex':   ['Male', 'Female', 'Male', 'Female', 'Male', 'Female'],
            'name':  ['Liam', 'Emma', 'Liam', 'Emma', 'Noah', 'Emma'],
            'count': [19837, 18688, 18798, 19800, 19117, 19496]}

filtered = filterTopNamesDoL(topnames, 19000)
assert len(filtered['count']) == 4
assert len(filtered['year']) == 4
assert len(filtered['name']) == 4
assert len(filtered['sex']) == 4

assert len(topnames['count']) == 6
assert len(topnames['year']) == 6
assert len(topnames['name']) == 6
assert len(topnames['sex']) == 6
```

**Q4** Write a function 

```
    addCatColumnDoL(tnDoL, threshold1, threshold2)
```

that adds a categorical column to a DoL representation in parameter `tnDoL` with the new column named `category` whose values are the strings `"small"` when count is below `threshold1`, is `"medium"` when count is greater than or equal to `threshold1` and less than `threshold2`, and `large` when count is greater than or equal to `threshold2`.

This change to `tnDoL` happens in place, rather than creating a new dictionary, and so nothing is returned from the function.

<!-- grade_id: cell-5a913843f1fafff9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Printing cell
topnames = {'year':  [2018, 2018, 2017, 2017, 2016, 2016],
            'sex':   ['Male', 'Female', 'Male', 'Female', 'Male', 'Female'],
            'name':  ['Liam', 'Emma', 'Liam', 'Emma', 'Noah', 'Emma'],
            'count': [19837, 18688, 18798, 19800, 19117, 19496]}

addCatColumnDoL(topnames, 19000, 19500)
print(topnames)
```

<!-- grade_id: cell-fec37caa6e87b5ef, points: 3, grade: true, solution: false -->
```python
# Testing cell
assert addCatColumnDoL.__doc__ is not None # don't forget the docstring!

topnames = {'year':  [2018, 2018, 2017, 2017, 2016, 2016],
            'sex':   ['Male', 'Female', 'Male', 'Female', 'Male', 'Female'],
            'name':  ['Liam', 'Emma', 'Liam', 'Emma', 'Noah', 'Emma'],
            'count': [19837, 18688, 18798, 19800, 19117, 19496]}

res = addCatColumnDoL(topnames, 19000, 19500)
assert res is None

assert 'category' in topnames
assert 'small' in topnames['category']
assert 'medium' in topnames['category']
assert 'large' in topnames['category']
assert len(topnames['category']) == len(topnames['year'])
```

> You've reached the final checkpoint in the lab.  Make sure to have it signed off by the instructor.
>
> Checkpoint 4: Let's imagine you added the `category` column to our `topnames` DoL, only to realize the second threshold was incorrect.  What would you have to do to update the DoL in place based on a new value of `threshold2`?

---

---
## Part G

How much time (in minutes/hours) did you spend on this lab outside of class?

YOUR ANSWER HERE