# Denison CS181/DA210 Homework 2.f - Step 2

Before you turn this problem in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells** (in the menubar, select Kernel$\rightarrow$Restart And Run All).

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import os
import os.path
import pandas as pd

datadir = "publicdata"
```

---

## Part B: Data normalization

**Q4:** Consider the file `ratings.csv`.  It has columns for first name, last name, `RatingA`, used for rating a particular restaurant (A), and `RatingB`, used for rating a different restaurant (B).

Load this dataset into a `DataFrame` called `ratings_untidy`.

<!-- grade_id: cell-453406ce5397fdb5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# Display the DataFrame
ratings_untidy
```

<!-- grade_id: cell-73c2bce4601a51de, points: 3, grade: true, solution: false -->
```python
# Testing cell
assert ratings_untidy.shape == (3,4)
assert ratings_untidy.iloc[0,1] == "Hirst"
assert list(ratings_untidy.iloc[:,2]) == [73, 57, 85]
assert list(ratings_untidy.iloc[:,3]) == [52, 72, 69]
```

**Q5:** This dataset is not tidy for the following reasons:
* The name of a "rater" should be a single variable (with values of the form "\<first\> \<last\>", like "Wei Chang").
* The particular restaurants are *values* of the data set.

Solve the first of these issues by combining the first two columns (for first and last name) into a single `Rater` column.  Name your resultant `DataFrame` `ratings_untidy2`.

Do not give the new data set a row label index.

```python
# YOUR CODE HERE
raise NotImplementedError()

# Display the DataFrame
ratings_untidy2
```

```python
# Testing cell
assert ratings_untidy2.shape == (3,3)
assert "First" not in ratings_untidy2
assert "Last" not in ratings_untidy2
assert "RatingA" in ratings_untidy2
assert "RatingB" in ratings_untidy2
assert "Rater" in ratings_untidy2
assert list(ratings_untidy2["Rater"]) == ["Hamid Hirst",
                                          "Scarlet Crossley",
                                          "Wei Chang"]
```

**Q6:** Finally, fix the second issue to transform `ratings_untidy2` into a tidy data set, naming it `ratings_tidy`.

Note that you should both transform the `RatingA` and `RatingB` columns, but your resulting column representing the restaurant variable, `Restaurant`, should have values `'A'` and `'B'` rather than `'RatingA'` and `'RatingB'`.  Your value column should be named `Rating`.

Once again, do not give the new data set a row label index.

```python
# YOUR CODE HERE
raise NotImplementedError()

# Display the tidy DataFrame
ratings_tidy
```

```python
# Testing cell
assert ratings_tidy.shape == (6,3)
assert list(ratings_tidy.columns) == ["Rater", "Restaurant", "Rating"]
assert list(ratings_tidy.loc[ratings_tidy["Restaurant"] == "A", "Rating"]) == [73,57,85]
assert list(ratings_tidy.loc[ratings_tidy["Restaurant"] == "B", "Rating"]) == [52,72,69]
```

---

---

## Part C

**Q7:** How much time (in minutes/hours) did you spend on this homework assignment?

YOUR ANSWER HERE

**Q8:** Who was your partner for this assignment?  If you worked alone, say so instead.

YOUR ANSWER HERE