# Denison CS181/DA210 Homework 2.f - Step 1

Before you turn this problem in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells** (in the menubar, select Kernel$\rightarrow$Restart And Run All).

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import os
import os.path
import pandas as pd

datadir = "publicdata"
```

---

## Part A: Transformation operations: transpose and melt

**Q1:** Make the following into a `pandas` data frame.  Assign it to `df`.
```
    {'A': {0: 'a', 1: 'b', 2: 'c'},
     'B': {0:  2,  1:  4,  2:  6},
     'C': {0:  1,  1:  3,  2:  5},
     'D': {0:  1,  1:  2,  2:  4}}
```

<!-- grade_id: cell-c93bf4fe3e267518, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# Display the DataFrame
df
```

<!-- grade_id: cell-886fdd47df433623, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert df.shape == (3,4)
assert list(df.columns) == ["A", "B", "C", "D"]
assert list(df.index) == list(range(3))
```

**Q2:** Perform a transpose operation on `df`, storing the result in `df_transpose`.

```python
# YOUR CODE HERE
raise NotImplementedError()

# Display the transposed DataFrame
df_transpose
```

```python
# Testing cell
assert df_transpose.shape == (4,3)
assert list(df_transpose.index) == ["A", "B", "C", "D"]
assert list(df_transpose.columns) == list(range(3))
```

```python
# Display the original DataFrame again
df
```

**Q3:** Suppose that we have determined that columns `B` and `D` are really _values_ of a _variable_ called `X`.  (This is similar to when columns represented values of the variable `Year` for `topnames`.)

We can make this dataset tidy using a `melt` operation.  Perform this operation and assign the result to `df_melted`.  The values of these columns should be placed in a new column named `Y`.

<!-- grade_id: cell-409fcf1eacadcf13, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# Display the DataFrame
df_melted
```

<!-- grade_id: cell-bc651d3241e7affb, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert df_melted.shape == (6,4)
assert "B" not in df_melted
assert "D" not in df_melted
assert "X" in df_melted
assert "Y" in df_melted
assert list(df_melted.loc[df_melted["X"] == "B", "A"]) == ['a','b','c']
assert list(df_melted.loc[df_melted["X"] == "B", "Y"]) == [2,4,6]
assert list(df_melted.loc[df_melted["X"] == "D", "Y"]) == [1,2,4]
```