# Denison CS181/DA210 SW Lab #1 - Step 3

Before you turn this problem in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells** (in the menubar, select Kernel$\rightarrow$Restart And Run All).

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import os
import os.path
import io
import sys
from contextlib import redirect_stdout

datadir = "./publicdata"
```

## Instructions

- Complete all the problems below by replacing the two-line sequence of `# YOUR CODE HERE` and `raise NotImplementedError()` with your solution code.
- All functions that you define for software labs **must have docstrings**.

**Q1** Write a function

```
    myFileList(directory)
```
    
that queries the specified directory, `directory`, for the files and folders contained therein and generates and returns a list with an element for each of the non-hidden **files** in the directory.  This means that your list should **not** include any entries for hidden files (that start with a `.` character), nor should it include entries for directories (i.e., folders).

<!-- grade_id: cell-5a2690ad91df7a9f, points: 0, grade: false, solution: true -->
```python
# Solution cell

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-07f7d7e0c1519cdf, points: 1, grade: true, solution: false -->
```python
# Testing cell
assert myFileList.__doc__ is not None # don't forget the docstring!

swlab_01_filelist = myFileList(".")
assert len(swlab_01_filelist) == 2
assert "filebasic.ipynb" in swlab_01_filelist
assert "swlab_01.md" in swlab_01_filelist
```

<!-- grade_id: cell-4379c9eebdfdaf7e, points: 2, grade: true, solution: false -->
```python
# Testing cell with tougher test cases

publicdata_filelist = myFileList("publicdata")
assert "tennyson.txt" in publicdata_filelist
```

> You've reached the final checkpoint in the lab.  Make sure to have it signed off by the instructor.
>
> Checkpoint 3: How many elements do you expect to be returned by `myFileList("publicdata")`?  Why?

---

---

## Time Spent

How much time (in minutes/hours) did you spend on this software lab (outside of class)?

YOUR ANSWER HERE