# Denison CS181/DA210 Homework 1.e - Step 2

Before you turn this problem in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells** (in the menubar, select Kernel$\rightarrow$Restart And Run All).

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import os
import os.path

datadir = "publicdata"
```

---

## Part C

**Q1** Assuming that `path` refers to a CSV file with one header line (with columns `year`, `sex`, `name`, `count`) and data lines with the same four fields, write a function

```
    readTopNamesLoL(path)
```

that reads the file and creates a LoL representation, returning both the list of column names and the list of lists structure from the function.

<!-- grade_id: cell-a2f3d39ffa88a59d, points: 0, grade: false, solution: true -->
```python
def readTopNamesLoL(path):
    '''
    Reads a CSV file and returns its column names and LoL representation
    '''
    f = open(path, 'r')
    headers = f.readline().strip().split(',')
    s = f.readlines()
    newLoL = []
    for line in s:
        line = line.strip().split(',')
        newlist = [int(item) if item.isdigit() else item for item in line]
        newLoL.append(newlist)
    return headers, newLoL
```

```python
# Debugging cell
tn10columns, tn10data = readTopNamesLoL(os.path.join(datadir, "tn10.csv"))
print(tn10columns)
print(tn10data)
```

Output:
```
['year', 'sex', 'name', 'count']
[[2014, 'Female', 'Emma', 20936], [2014, 'Male', 'Noah', 19305], [2015, 'Female', 'Emma', 20455], [2015, 'Male', 'Noah', 19635], [2016, 'Female', 'Emma', 19496], [2016, 'Male', 'Noah', 19117], [2017, 'Female', 'Emma', 19800], [2017, 'Male', 'Liam', 18798], [2018, 'Female', 'Emma', 18688], [2018, 'Male', 'Liam', 19837]]
```

<!-- grade_id: cell-06e29c14178a4333, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert readTopNamesLoL.__doc__ is not None # don't forget the docstring!

tn10columns, tn10data = readTopNamesLoL(os.path.join(datadir, "tn10.csv"))

assert isinstance(tn10data, list)
assert len(tn10data) == 10

assert 'year' in tn10columns
assert 'sex' in tn10columns
assert 'count' in tn10columns
assert 'name' in tn10columns

assert len(tn10data[0]) == 4
assert tn10data[0] == [2014, "Female", "Emma", 20936]
```

**Q2** Write a function 

```
    filterTopNamesLoL(tnLoL, threshold)
```

to create a **filter** copy of a topnames LoL `tnLoL` (with columns `year`, `sex`, `name`, `count`) so that only rows with a count value greater than or equal to `threshold` are present in the newly created LoL.

Note that you are creating a **new** LoL with the filtered data, and **not** modifying `tnLoL` in place.  Your function should return the new LoL.

<!-- grade_id: cell-a97c99c74a286d53, points: 0, grade: false, solution: true -->
```python
filterTopNamesLoL = lambda tnLoL, threshold : [item for item in tnLoL if item[3] >= threshold]
filterTopNamesLoL.__doc__ = '''Returns a new LoL with names whose count are higher than threshold'''
```

```python
# Debugging cell
tn10columns, tn10data = readTopNamesLoL(os.path.join(datadir, "tn10.csv"))
tncolumns,tndata = readTopNamesLoL(os.path.join(datadir, "topnames.csv"))
tn_filter = filterTopNamesLoL(tn10data, 19000)
print(tn_filter)
```

Output:
```
[[2014, 'Female', 'Emma', 20936], [2014, 'Male', 'Noah', 19305], [2015, 'Female', 'Emma', 20455], [2015, 'Male', 'Noah', 19635], [2016, 'Female', 'Emma', 19496], [2016, 'Male', 'Noah', 19117], [2017, 'Female', 'Emma', 19800], [2018, 'Male', 'Liam', 19837]]
```

<!-- grade_id: cell-e47341b9e033c921, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert filterTopNamesLoL.__doc__ is not None # don't forget the docstring!

topnames = [[2018, 'Male', 'Liam', 19837],
            [2018, 'Female', 'Emma', 18688],
            [2017, 'Male', 'Liam', 18798],
            [2017, 'Female', 'Emma', 19800],
            [2016, 'Male', 'Noah', 19117],
            [2016, 'Female', 'Emma', 19496]]
columns = ['year', 'sex', 'name', 'count']

filtered = filterTopNamesLoL(topnames, 19000)
assert len(filtered) == 4
```

**Q3** Write a function 

```
    addCatColumnLoL(tnLoL, threshold1, threshold2)
```

that adds a categorical column to a LoL representation in parameter `tnLoL` with the new column named `category` whose values are the strings `"small"` when count is below `threshold1`, is `"medium"` when `count` is greater than or equal to `threshold1` and less than `threshold2`, and `large` when count is greater than or equal to `threshold2`.

This change to `tnLoL` happens in place, rather than creating a new list of lists, and so nothing is returned from the function.

```python
def addCatColumnLoL(tnLoL, threshold1, threshold2):
    '''
    Categorize the list of names according to count based on 2 thresholds and update existing LoL accordingly
    '''
    for item in tnLoL:
        if item[3] < threshold1:
            item.append('small')
        elif item[3] >= threshold1 and item[3] < threshold2:
            item.append('medium')
        else:
            item.append('large')
```

```python
# Printing cell
topnames = [[2018, 'Male', 'Liam', 19837],
            [2018, 'Female', 'Emma', 18688],
            [2017, 'Male', 'Liam', 18798],
            [2017, 'Female', 'Emma', 19800],
            [2016, 'Male', 'Noah', 19117],
            [2016, 'Female', 'Emma', 19496]]

addCatColumnLoL(topnames, 19000, 19500)
print(topnames)
```

Output:
```
[[2018, 'Male', 'Liam', 19837, 'large'], [2018, 'Female', 'Emma', 18688, 'small'], [2017, 'Male', 'Liam', 18798, 'small'], [2017, 'Female', 'Emma', 19800, 'large'], [2016, 'Male', 'Noah', 19117, 'medium'], [2016, 'Female', 'Emma', 19496, 'medium']]
```

```python
# Testing cell
assert addCatColumnLoL.__doc__ is not None # don't forget the docstring!

topnames = [[2018, 'Male', 'Liam', 19837],
            [2018, 'Female', 'Emma', 18688],
            [2017, 'Male', 'Liam', 18798],
            [2017, 'Female', 'Emma', 19800],
            [2016, 'Male', 'Noah', 19117],
            [2016, 'Female', 'Emma', 19496]]

res = addCatColumnLoL(topnames, 19000, 19500)
assert res is None

assert len(topnames) == 6
assert isinstance(topnames[0], list)
assert len(topnames[0]) == 5
assert topnames[0][4] == "large"
assert topnames[2][4] == "small"
assert topnames[5][4] == "medium"
```

---

---
## Part D

How much time (in minutes/hours) did you spend on this homework assignment?

30 minutes