# Denison CS181/DA210 Homework 1.c - Step 1

Before you turn this problem in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells** (in the menubar, select Kernel$\rightarrow$Restart And Run All).

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

Also, note that all functions that you define for homework **must have docstrings**.

---

```python
import os
import os.path
import io
import sys
from contextlib import redirect_stdout
```

---
## Part A

**Q1** Write a function

```
    decreaseList(L, amount)
```

that, for each number in list L, decreases L by `amount`.  The list should be modified in place, and the function should not create a new list, nor return any value.

<!-- grade_id: cell-2370a7c1bbe407fc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-012b8ea5dfc08243, points: 2, grade: true, solution: false -->
```python
# Testing cell
L = [0.65, 0.75, 0.85]
decreaseList(L, 0.01)
assert L == [0.64, 0.74, 0.84]

data = [3,4,5]
decreaseList(amount = 1, L = data)
assert data == [2,3,4]
```

**Q2** Write a function `swap_case(s)` that swap cases in the string `s`. In other words, convert all lowercase letters to uppercase letters and vice versa. Do not use the `swapcase()` method. Instead, use `ord()` and `chr()`. For example, if `s = "My string"` you return `"mY STRING"`.

<!-- grade_id: cell-9e32228bd18a25d0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3b660b9b6f4b7390, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert swap_case("WikiPedia") == "wIKIpEDIA"
assert swap_case("Lorem Ipsum") == "lOREM iPSUM"
assert swap_case("The quick Brown fOX jumped Over THE laZy doG") == "tHE QUICK bROWN Fox JUMPED oVER the LAzY DOg"
```

---
## Part B

**Q3** Implement a function

```
    insertElement(myList, e, i)
```
 
that inserts an element `e` at index `i` of `myList` in place.

<!-- grade_id: cell-453c1aac75de2844, points: 0, grade: false, solution: true -->
```python
def insertElement(myList, e, i):
    """
    Inserts an element e into index i of myList.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b8fd5f38b0641496, points: 2, grade: true, solution: false -->
```python
# Testing cell
myList =[1, 2, 3, 4]
assert insertElement(myList, 5, 0) == None
assert myList[0] == 5
insertElement(myList, 6, 1)
assert myList[0] == 5
assert myList[1] == 6
insertElement(myList, 7, len(myList))
assert len(myList) == 7
assert myList[-1] == 7
```

**Q4** Implement a function `findElement(myList, e)` that returns the index of the first occurence of element `e` in `myList`. Do not use any built-in functions.

<!-- grade_id: cell-061721e509002e48, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d5fd301307dccd31, points: 2, grade: true, solution: false -->
```python
# Testing cell
myList = [3, 7]
assert findElement(myList, 7) == 1
assert findElement(myList, 14) == None
```