# Denison CS181/DA210 Homework 2.e - Step 1

Before you turn this problem in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells** (in the menubar, select Kernel$\rightarrow$Restart And Run All).

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import os
import os.path
import pandas as pd

datadir = "publicdata"
```

---

## Part A: Concatenating along the row dimension

**Q1:** In the data directory, you will find two csv files, `educationTop.csv` and `educationBottom.csv`, both based on data hosted by `www.census.gov`.  Both have the same columns, and each row is a U.S. metropolitan area, with data on population, education, and unemployment.  The first `csv` file contains metropolitan areas starting with A-K, and the second starting with L-Z.

Read both into `pandas` data frames `education_top` and `education_bottom`, using the column `GEO.id2` as an index.

Be careful, as these files have **two** lines that might be considered header lines.  The first one is more descriptive, and the second one is more concise and gives succint variable names.  So when we read from the csv, we want to **skip** the first line and then use the second line as the header line.  The `pandas` `read_csv` constructor can do this for you, but you have to discover and use the correct arguments.  (Documentation: https://pandas.pydata.org/docs/reference/api/pandas.read_csv.html)

<!-- grade_id: cell-b25fab8c12339430, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Display a sample of the "top" DataFrame
education_top.head(2)
```

```python
# Display a sample of the "bottom" DataFrame
education_bottom.head(2)
```

<!-- grade_id: cell-8f07ee547e7c1ff8, points: 3, grade: true, solution: false -->
```python
# Testing Cell
assert education_top.shape == (176, 60)
assert education_bottom.shape == (190, 60)
assert education_top.index.name == "GEO.id2"
assert education_bottom.index.name == "GEO.id2"
assert "HD01_VD01" in education_top
assert "HD01_VD01" in education_bottom
assert education_top.iloc[8,2] == 129504
assert education_bottom.iloc[8,4] == 39259
```

**Q2:** Concatenate these two data frames along the row dimension (with the top one on top), and call the result `educationDF`.

```python
# YOUR CODE HERE
raise NotImplementedError()

# Display a sample of the DataFrame
educationDF.head(4)
```

```python
# Testing Cell
assert educationDF.shape == (366, 60)
assert educationDF.index.name == "GEO.id2"
assert "HD01_VD01" in educationDF
assert educationDF.iloc[8,2] == 129504
assert educationDF.iloc[184,4] == 39259
```

---

## Part B: Concatenating along the column dimension

**Q3:** In the data directory, you will find two csv files, `educationLeft.csv` and `educationRight.csv`, both based on data hosted by `www.census.gov`.  Both have the same rows, and each row is a U.S. metropolitan area, with data on population, education, and unemployment.  The first has information on individuals without a college degree, and the second has information on individuals with a college degree.

Read both into `pandas` data frames `education_left` and `education_right`.  The `GEO.display-label` field should serve as the row index.

The same cautions on the CSV having an informational row of headers followed by the header variable names applies to these two CSV files as well.

<!-- grade_id: cell-698c9ada24032095, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Display a sample of the left DataFrame
education_left.head(2)
```

```python
# Display a sample of the right DataFrame
education_right.head(2)
```

<!-- grade_id: cell-f413a9fabbcd0427, points: 3, grade: true, solution: false -->
```python
# Testing Cell
assert education_left.shape == (366, 30)
assert education_right.shape == (366, 28)
assert education_left.index.name == "GEO.display-label"
assert education_right.index.name == "GEO.display-label"
assert "HD01_VD01" in education_left
assert "HD01_VD01" not in education_right
assert "Abilene, TX Metro Area" in education_left.index
assert "Abilene, TX Metro Area" in education_right.index
assert education_left.iloc[10,3] == 2265
assert education_right.iloc[10,2] == 63125
```

**Q4:** Concatenate these two data frames along the column dimension, and call the result `educationDF2`.

```python
# YOUR CODE HERE
raise NotImplementedError()

# Display a sample of the DataFrame
educationDF2.head(4)
```

```python
# Testing Cell
assert educationDF2.shape == (366, 58)
assert educationDF2.index.name == "GEO.display-label"
assert "HD01_VD01" in educationDF2
assert educationDF2.iloc[10,3] == 2265
assert educationDF2.iloc[10,32] == 63125
```

---

## Part C: Joining/merging data frames

**Q5:** In the data directory, you will find two csv files, `educationLeftJ.csv` and `educationRightJ.csv`, both based on data hosted by `www.census.gov`.  In both, rows represent U.S. metropolitan area, with data on population, education, and unemployment.  However, they do not have exactly the same set of rows, and the columns are totally different except for the index column `Geography`.

In these CSV files, there is not an "extra" header line.  The row label index should come from `Geography` unless you are instructed otherwise.

Read both into `pandas` data frames, with names `education_left_j` and `education_right_j`.

<!-- grade_id: cell-f3cb3f384a53b7b2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Display a sample of the left_j DataFrame
education_left_j.head(2)
```

```python
# Display a sample of the right_j DataFrame
education_right_j.head(2)
```

<!-- grade_id: cell-a1e03d05287e5c02, points: 4, grade: true, solution: false -->
```python
# Testing Cell
assert education_left_j.shape == (249, 12)
assert education_right_j.shape == (215, 12)
assert education_left_j.index.name == "Geography"
assert education_right_j.index.name == "Geography"
assert "Estimate; Less than high school graduate: - In labor force: - Civilian: - Unemployed" in education_left_j
assert "Estimate; Less than high school graduate: - In labor force: - Civilian: - Unemployed" not in education_right_j
assert "Estimate; Bachelor's degree or higher: - In labor force:" not in education_left_j
assert "Estimate; Bachelor's degree or higher: - In labor force:" in education_right_j
assert "Battle Creek, MI Metro Area" in education_left_j.index
assert "Battle Creek, MI Metro Area" not in education_right_j.index
assert "Bridgeport-Stamford-Norwalk, CT Metro Area" not in education_left_j.index
assert "Bridgeport-Stamford-Norwalk, CT Metro Area" in education_right_j.index
assert education_left_j.iloc[5,6] == 38994
assert education_right_j.iloc[9,2] == 59784
```

**Q6:** Make a copy of `education_left_j` called `education_left_orig`.  Starting with `education_left_j`, do a left join to bring in the data from `education_right_j`, storing your answer as `df_j`.

```python
# YOUR CODE HERE
raise NotImplementedError()

# Display a sample of the joined DataFrame
df_j1.head(2)
```

```python
# Testing Cell
assert education_left_orig.shape == (249, 12)

assert df_j1.shape == (249, 24)
assert df_j1.index.name == "Geography"
assert "Estimate; Less than high school graduate: - In labor force: - Civilian: - Unemployed" in df_j1
assert "Estimate; Bachelor's degree or higher: - In labor force:" in df_j1
assert "Battle Creek, MI Metro Area" in df_j1.index
assert "Bridgeport-Stamford-Norwalk, CT Metro Area" not in df_j1.index
assert df_j1.iloc[5,6] == 38994
assert df_j1.iloc[5,14] == 59784
```

**Q7:** Starting with `education_left_orig`, do an inner join to bring in the data from `education_right_j`, storing your answer as `df_j2`.

```python
# YOUR CODE HERE
raise NotImplementedError()

# Display a sample of the joined DataFrame
df_j2.head(2)
```

```python
# Testing Cell
assert df_j2.shape == (152, 24)
assert df_j2.index.name == "Geography"
assert "Estimate; Less than high school graduate: - In labor force: - Civilian: - Unemployed" in df_j2
assert "Estimate; Bachelor's degree or higher: - In labor force:" in df_j2
assert "Battle Creek, MI Metro Area" not in df_j2.index
assert "Bridgeport-Stamford-Norwalk, CT Metro Area" not in df_j2.index
assert "Asheville, NC Metro Area" in df_j2.index
assert df_j2.iloc[5,6] == 38994
assert df_j2.iloc[5,14] == 59784
```

**Q8:** We can "merge" two data frames if they do not have a meaningful index, but do share a column name.  You can use this syntax:

```
    merged_df = pd.merge(df1, df2, on=<colname>, how=<merge_type>)
```

where `<colname>` is the name of the column they share (typically a string), and `<merge_type>` is a string like `"left"`, `"inner"`, etc.  We didn't get to this in class, but you can see examples in Section 8.3.3.2 of the book.

Now read the original csv files in as `edu_left` and `edu_right` with no meaningful index. Then, starting from `edu_left`, do an inner merge along the column `Geography`, storing your answer as `df_j3`.

```python
# YOUR CODE HERE
raise NotImplementedError()

# Display a sample of the merged DataFrame
df_j3.head(2)
```

```python
# Testing Cell
assert df_j3.shape == (152, 25)
assert df_j3.index.name == None
assert "Estimate; Less than high school graduate: - In labor force: - Civilian: - Unemployed" in df_j3
assert "Estimate; Bachelor's degree or higher: - In labor force:" in df_j3
assert "Battle Creek, MI Metro Area" not in list(df_j3.Geography)
assert "Bridgeport-Stamford-Norwalk, CT Metro Area" not in list(df_j3.Geography)
assert "Asheville, NC Metro Area" in list(df_j3.Geography)
assert df_j3.iloc[5,7] == 38994
assert df_j3.iloc[5,15] == 59784
```

**Q9:** Explain the difference in the number of rows and columns in `df_j1`, `df_j2`, and `df_j3`.  Be specific.

<!-- grade_id: cell-60bd9e8dd53dd5f0, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

---

---

## Part D

**Q10:** How much time (in minutes/hours) did you spend on this homework assignment?

YOUR ANSWER HERE

**Q11:** Who was your partner for this assignment?  If you worked alone, say so instead.

YOUR ANSWER HERE