# Denison CS181/DA210 SW Lab #3 - Step 1

Before you turn this problem in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells** (in the menubar, select Kernel$\rightarrow$Restart And Run All).

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import os
import os.path

datadir = "publicdata"
```

---

## Part A: List of Dictionaries (LoD)

In this representation, each row is represented by an individual dictionary.  Thus, each entry in the LoD represents **exactly one row** of the data set.

Each row dictionary has one key-value pair for each column in the dataset.  For example, in `topnames.csv`, the column names (and thus keys) are `"year"`, `"sex"`, `"name"`, and `"count"`.

Here is an example hard-coded LoD corresponding to a snippet of `topnames.csv`:

```python
LoD = [{'year': 2014, 'sex': 'Female', 'name': 'Emma', 'count': 20936}, 
       {'year': 2014, 'sex': 'Male', 'name': 'Noah', 'count': 19305}, 
       {'year': 2015, 'sex': 'Female', 'name': 'Emma', 'count': 20455}, 
       {'year': 2015, 'sex': 'Male', 'name': 'Noah', 'count': 19635}, 
       {'year': 2016, 'sex': 'Female', 'name': 'Emma', 'count': 19496}, 
       {'year': 2016, 'sex': 'Male', 'name': 'Noah', 'count': 19117}, 
       {'year': 2017, 'sex': 'Female', 'name': 'Emma', 'count': 19800}, 
       {'year': 2017, 'sex': 'Male', 'name': 'Liam', 'count': 18798}, 
       {'year': 2018, 'sex': 'Female', 'name': 'Emma', 'count': 18688}, 
       {'year': 2018, 'sex': 'Male', 'name': 'Liam', 'count': 19837}]
```

When a representation makes it easy to work with entire rows, but working with columns is more cumbersome, it is a **row-centric** representation.

Here is how you can get a row of data with a LoD:

```python
# Row operation: all values of row with index 4
row_index = 3
print(LoD[row_index])
```

It is much more work to get a column of data with a LoD:

```python
# Column operation: all values of 'count' column
print(LoD[0]["count"], LoD[1]["count"], LoD[2]["count"], LoD[3]["count"], LoD[4]["count"], LoD[5]["count"], LoD[6]["count"], LoD[7]["count"], LoD[8]["count"], LoD[9]["count"])
```

Fortunately, list comprehensions make column operations for a LoD a little easier:

```python
# Column operation: all values of 'count' column using a list comprehension
print([row["count"] for row in LoD])
```

> You've reached the first checkpoint in the lab.  Make sure to have it signed off by the instructor or TA.
>
> Checkpoint 1: What expression could you use to get the number of rows in `LoD`.  What happens to the LoD if a row is missing the data for one of the columns?

---

## Part B

**Q1** Assume that `path` refers to a CSV file with the same format used in our variations of `topnames` with a header line containing the four column names, and then $x$ rows of comma-separated field values.  Write a function

```
    readTopNamesLoD(path)
```

that reads the file and creates a LoD representation and returns the structure.

<!-- grade_id: cell-3ac611d5fbae072d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Debugging cell
tn10data = readTopNamesLoD(os.path.join(datadir, "tn10.csv"))
print(tn10data)
```

<!-- grade_id: cell-ab5935ca57c293ca, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert readTopNamesLoD.__doc__ is not None # don't forget the docstring!

tn10data = readTopNamesLoD(os.path.join(datadir, "tn10.csv"))

assert isinstance(tn10data, list)
assert len(tn10data) == 10

assert len(tn10data[0]) == 4
assert isinstance(tn10data[0], dict)

assert tn10data[0]['sex'] == 'Female'
assert tn10data[0]['year'] == 2014
assert tn10data[0]['name'] == 'Emma'
assert tn10data[0]['count'] == 20936
```

**Q2** Write a function 

```
    filterTopNamesLoD(tnLoD, threshold)
```

to create a **filter** copy of a topnames LoD `tnLoD` (with columns `year`, `sex`, `name`, `count`) so that only rows with a `count` value greater than or equal to `threshold` are present in the newly created LoD.

Your function should return the new and filtered LoD, nd not modify the original `tnLoD`.

```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Debugging cell
tn10 = readTopNamesLoD(os.path.join(datadir, "tn10.csv"))
tn10_filter = filterTopNamesLoD(tn10, 19000)
print(tn10_filter)
```

```python
# Testing cell
assert filterTopNamesLoD.__doc__ is not None # don't forget the docstring!

LoD = [{'year': 2014, 'sex': 'Female', 'name': 'Emma', 'count': 20936}, 
       {'year': 2014, 'sex': 'Male', 'name': 'Noah', 'count': 19305}, 
       {'year': 2015, 'sex': 'Female', 'name': 'Emma', 'count': 20455}, 
       {'year': 2015, 'sex': 'Male', 'name': 'Noah', 'count': 19635}, 
       {'year': 2016, 'sex': 'Female', 'name': 'Emma', 'count': 19496}, 
       {'year': 2016, 'sex': 'Male', 'name': 'Noah', 'count': 19117}, 
       {'year': 2017, 'sex': 'Female', 'name': 'Emma', 'count': 19800}, 
       {'year': 2017, 'sex': 'Male', 'name': 'Liam', 'count': 18798}, 
       {'year': 2018, 'sex': 'Female', 'name': 'Emma', 'count': 18688}, 
       {'year': 2018, 'sex': 'Male', 'name': 'Liam', 'count': 19837}]

filtered = filterTopNamesLoD(LoD, 19000)
assert len(filtered) == 8
assert len(filtered[0]) == 4

assert len(LoD) == 10
assert len(LoD[0]) == 4
```

**Q3** Write a function 

```
    addCatColumnLoD(tnLoD, threshold1, threshold2)
```

that adds a categorical column to a LoD representation in parameter `tnLoD` with the new column named `category` whose values are the strings `"small"` when count is below `threshold1`, is `"medium"` when count is greater than or equal to `threshold1` and less than `threshold2`, and `"large"` when count is greater than or equal to `threshold2`.

This change to tnLoD happens in place, rather than creating a new list, and so nothing is returned from the function.

<!-- grade_id: cell-0d9878d3fcbb29fa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Debugging cell
LoD = [{'year': 2014, 'sex': 'Female', 'name': 'Emma', 'count': 20936}, 
       {'year': 2014, 'sex': 'Male', 'name': 'Noah', 'count': 19305}, 
       {'year': 2015, 'sex': 'Female', 'name': 'Emma', 'count': 20455}, 
       {'year': 2015, 'sex': 'Male', 'name': 'Noah', 'count': 19635}, 
       {'year': 2016, 'sex': 'Female', 'name': 'Emma', 'count': 19496}, 
       {'year': 2016, 'sex': 'Male', 'name': 'Noah', 'count': 19117}, 
       {'year': 2017, 'sex': 'Female', 'name': 'Emma', 'count': 19800}, 
       {'year': 2017, 'sex': 'Male', 'name': 'Liam', 'count': 18798}, 
       {'year': 2018, 'sex': 'Female', 'name': 'Emma', 'count': 18688}, 
       {'year': 2018, 'sex': 'Male', 'name': 'Liam', 'count': 19837}]

addCatColumnLoD(LoD, 19000, 19500)
print(LoD)
```

<!-- grade_id: cell-1da7403c4d9cd8aa, points: 3, grade: true, solution: false -->
```python
# Testing cell
assert addCatColumnLoD.__doc__ is not None # don't forget the docstring!

LoD = [{'year': 2014, 'sex': 'Female', 'name': 'Emma', 'count': 20936}, 
       {'year': 2014, 'sex': 'Male', 'name': 'Noah', 'count': 19305}, 
       {'year': 2015, 'sex': 'Female', 'name': 'Emma', 'count': 20455}, 
       {'year': 2015, 'sex': 'Male', 'name': 'Noah', 'count': 19635}, 
       {'year': 2016, 'sex': 'Female', 'name': 'Emma', 'count': 19496}, 
       {'year': 2016, 'sex': 'Male', 'name': 'Noah', 'count': 19117}, 
       {'year': 2017, 'sex': 'Female', 'name': 'Emma', 'count': 19800}, 
       {'year': 2017, 'sex': 'Male', 'name': 'Liam', 'count': 18798}, 
       {'year': 2018, 'sex': 'Female', 'name': 'Emma', 'count': 18688}, 
       {'year': 2018, 'sex': 'Male', 'name': 'Liam', 'count': 19837}]

res = addCatColumnLoD(LoD, 19000, 19500)
assert res is None

assert len(LoD) == 10
assert isinstance(LoD[0], dict)
assert len(LoD[0]) == 5
assert 'category' in LoD[0]
assert [ rowD['category'] for rowD in LoD ] == ['large', 'medium', 'large',
 'large', 'medium', 'medium', 'large', 'small', 'small', 'large']
```

> You've reached the second and final checkpoint in the lab.  Make sure to have it signed off by the instructor or TA.
>
> Checkpoint 2: Let's imagine you added the `category` column to our topnames LoD, only to realize the second threshold was incorrect.  What would you have to do to update the LoD in place based on a new value of `threshold2`?

---

---
## Part C

How much time (in minutes/hours) did you spend on this lab outside of class?

YOUR ANSWER HERE