# Denison CS181/DA210 Homework 3.b - Step 1

Before you turn this problem in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells** (in the menubar, select Kernel$\rightarrow$Restart And Run All).

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import io
from lxml import etree
import json
import sys
import os.path
import pandas as pd

datadir = "publicdata"
```

---

## Part A

**Q1:** Consider the following table of `subjects` data:

subject | name | department
---------|-----|--------------
CS | Computer Science | CS
MATH | Mathematics | MATH
ENGL | English Literature | ENGL

Using a *text editor*, edit and create a file named `subjects.xml` in the `datadir` directory (`"publicdata"`) that contains a legal XML representation of this data.

Once created, write a Python code sequence to read and parse the file, and then, using `etree.tostring()`, create a Python string to reference the (pretty-printed) decoded string version of the tree.  Name this variable `subjects_str`.

<!-- grade_id: cell-352a2a50e91b618c, points: 0, grade: false, solution: true -->
```python
tree = etree.parse(os.path.join(datadir, "subjects.xml"))
subjects_str = etree.tostring(tree, pretty_print = True).decode('utf-8')

# Pretty-print the XML data
print(subjects_str)
```

Output:
```
<subjects>
    <subject subject="CS" name="Computer Science" department="CS"/>
    <subject subject="MATH" name="Mathematics" department="MATH"/>
    <subject subject="ENGL" name="English Literature" department="ENGL"/>
</subjects>
```

<!-- grade_id: cell-33dd26f22f1d5e39, points: 3, grade: true, solution: false -->
```python
# Testing Cell
path = os.path.join(datadir, "subjects.xml")
assert os.path.isfile(path)
assert isinstance(subjects_str, str)
assert 75 < len(subjects_str)
```

**Q2:** Now consider the `courses` table below.  Using a text editor, edit and create `courses.xml` in the `datadir` directory (`"publicdata"`) that contains a legal XML tree representing this table:

subject | coursenum | title
---------|-----------|-----------------------
CS | 173 | Intermediate Computer Science
CS | 372 | Operating Systems
MATH | 300 | Introduction to Proofs
ENGL | 213 | Early British Literature

Once created, write a Python code sequence to read and parse the file, and then, as in Q1, use `etree.tostring()` to assign the string version of the courses tree to the variable `courses_str`.

<!-- grade_id: cell-90c8b3ec459f1a09, points: 0, grade: false, solution: true -->
```python
tree = etree.parse(os.path.join(datadir, "courses.xml"))
courses_str = etree.tostring(tree, pretty_print = True).decode('utf-8')

# Pretty-print the XML data
print(courses_str)
```

Output:
```
<courses>
    <course course="CS">
        <coursedata coursenum="173">
            <coursename title="Intermediate Computer Science"/>
        </coursedata>
        <coursedata coursenum="372">
            <coursename title="Operating Systems"/>
        </coursedata>
    </course>
    <course course="MATH">
        <coursedata coursenum="300">
            <coursename title="Introduction to Proofs"/>
        </coursedata>
    </course>
    <course course="ENGL">
        <coursedata coursenum="213">
            <coursename title="Early British Literature"/>
        </coursedata>
    </course>
</courses>
```

<!-- grade_id: cell-559f37743a566fde, points: 3, grade: true, solution: false -->
```python
# Testing Cell
path = os.path.join(datadir, "courses.xml")
assert os.path.isfile(path)
assert isinstance(courses_str, str)
assert 75 < len(courses_str)
```

**Q3:** Suppose you wanted a tree that contained **both** of the above tables.  Write a file named `school.xml` in the current directory that composes **as a single tree** both of the above component tables.

As before, once created, write a Python code sequence to read and parse the file.  Call the resulting string variable `school_str`.

<!-- grade_id: cell-04bdc0fe17d48c37, points: 3, grade: true, solution: true -->
```python
tree = etree.parse(os.path.join(datadir, "school.xml"))
school_str = etree.tostring(tree, pretty_print = True).decode('utf-8')

# Pretty-print the XML data
print(school_str)
```

Output:
```
<courses>
    <course course="CS" subject="CS" name="Computer Science" department="CS">
        <coursedata coursenum="173">
            <coursename title="Intermediate Computer Science"/>
        </coursedata>
        <coursedata coursenum="372">
            <coursename title="Operating Systems"/>
        </coursedata>
    </course>
    <course course="MATH" subject="MATH" name="Mathematics" department="MATH">
        <coursedata coursenum="300">
            <coursename title="Introduction to Proofs"/>
        </coursedata>
    </course>
    <course course="ENGL" subject="ENGL" name="English Literature" department="ENGL">
        <coursedata coursenum="213">
            <coursename title="Early British Literature"/>
        </coursedata>
    </course>
</courses>
```

**Q4:** Write a function:
```
    getLocalXML(filename, datadir=".", parser=None)
``` 
that performs the common steps of creating a path from the given `filename` and `datadir` and parses the XML file, using the passed `parser`, if any, and returns the Element at the **root** of the tree.  If `parser` is not passed, the standard `XMLParser` should be used.

If the file is not found, or if the parse is unsuccessful (due to XML not being "well formed"), the function should return `None`. Remember that if a parse is unsuccessful, the `etree` module raises an exception.  That means that you should have a `try` block, and indented within that block, the `parse()` invocation should occur.  The `try` block is followed by an `except Exception as e:` line, and within that, your return `None`.  If no exception is raised, code execution will proceed beyond the `try`/`except` block, and that is where you would return the root of the parsed tree.

<!-- grade_id: cell-81428647baf27b6c, points: 0, grade: false, solution: true -->
```python
def getLocalXML(filename, datadir=".", parser=None):
    '''
    Parse an XML file and returns Elements at the root of tree
    '''
    file_path = os.path.join(datadir, filename)
    try:
        tree0 = etree.parse(file_path)
        root0 = tree0.getroot()
    except Exception as e:
        return None
    return root0
```

<!-- grade_id: cell-6f7ca02f25f931f2, points: 3, grade: true, solution: false -->
```python
# Testing cell
assert getLocalXML.__doc__ is not None # don't forget the docstring!

myparser = etree.XMLParser(remove_blank_text=True)

wroot = getLocalXML("widombooks.xml", datadir, myparser)
assert len(wroot) == 8
bad = getLocalXML("foo.xml", datadir, myparser)
assert bad == None
bad2 = getLocalXML("bad.xml", datadir)
assert bad2==None
```

---

---

## Part B

**Q5:** How much time (in minutes/hours) did you spend on this homework assignment?

1.5 hours

**Q6:** Who was your partner for this assignment?  If you worked alone, say so instead.

I worked alone