# Denison CS181/DA210 Homework 1.b - Step 2

Before you turn this problem in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells** (in the menubar, select Kernel$\rightarrow$Restart And Run All).

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import os
import os.path
import io
import sys
from contextlib import redirect_stdout
```

## Instructions

- Complete all the problems below by replacing the two-line sequence of `# YOUR CODE HERE` and `raise NotImplementedError()` with your solution code.
- All functions that you define for homework **must have docstrings**.

---
## Part A

**Q1** Write a function that returns the square of the number $n$, if $n\geq 1$, and returns 0 otherwise. 

<!-- grade_id: cell-bfb578e126c1b55f, points: 0, grade: false, solution: true -->
```python
# Solution cell
def square(n):
    """
    Compute the square of the number n, or return 0 if less than 1.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c6379a4803420748, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert square.__doc__ is not None # don't remove the docstring!

"""Check that square returns the correct output for several inputs."""
assert square(10) == 100
assert square(1) == 1
assert square(2) == 4
assert square(11) == 121

"""Check that square returns None for invalid inputs."""
assert square(-1) == 0
```

**Q2** Using your `square` function, write a function that computes the sum of the squares of the numbers from 1 to $n$. Your function should call the `square` function -- it should NOT reimplement its functionality.

<!-- grade_id: cell-18bc2f28fbcc3eba, points: 0, grade: false, solution: true -->
```python
def sum_of_squares(n):
    """
    Compute the sum of the squares of numbers from 1 to n.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1706850a8a0c3286, points: 1, grade: true, solution: false -->
```python
"""Check that sum_of_squares returns the correct answer for various inputs."""
assert sum_of_squares(1) == 1
assert sum_of_squares(2) == 5
assert sum_of_squares(10) == 385
assert sum_of_squares(11) == 506
```

<!-- grade_id: cell-4db891d69ac65ac4, points: 1, grade: true, solution: false -->
```python
"""Check that sum_of_squares relies on square."""
orig_square = square
del square
try:
    sum_of_squares(1)
except NameError:
    pass
else:
    raise AssertionError("sum_of_squares does not use square")
finally:
    square = orig_square
```

**Q3** Write a function `sum_of_evens` that only adds up the even numbers between 0 and the number $n$.

<!-- grade_id: cell-9416b535d19fc949, points: 0, grade: false, solution: true -->
```python
def sum_of_evens(n):
    """
    Compute the sum of the even numbers from 0 to n.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d9c5a7f8882a8e91, points: 1, grade: true, solution: false -->
```python
"""Check that sum_of_evens returns the correct answer for various inputs."""
assert sum_of_evens(1) == 0
assert sum_of_evens(5) == 6
assert sum_of_evens(10) == 30
assert sum_of_evens(11) == 30
```

**Q4** Describe a **different** way to implement `sum_of_evens()`.

<!-- grade_id: cell-0083b8d23a14fc37, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q5** Use three *different* ways to create three lists with the integer numbers 1 through 10, inclusive.  Use the variable names `myList1`, `myList2`, and `myList3` to hold the results.

<!-- grade_id: cell-38d713dddb2b5cfe, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Q6** Write a function:

```
    mean(data)
```
 
that takes a list of numbers given by `data` and computes and returns the arithmetic mean of the values in the list.  Your function should work properly even if `data` is an empty list, and should return 0 in that case.

<!-- grade_id: cell-a32002531e280bd3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ef96cd921bb7b67b, points: 2, grade: true, solution: false -->
```python
# Testing cell
list1 = [3, 4, 7, 2]
assert mean(list1) == 4
list2 = []
assert mean(list2) == 0
```

---
## Part B

**Q7** For the following questions, each answer should consist of a single
`print` with a single argument--an expression involving the list `data`, which, for these questions, references the list `[32, 42, 11, 15, 58, 44, 16]`.

Leave the code we give you "as is", and your print should be indented so that your solution becomes the body of the `with`.

Print a list that includes all elements *except* the first element.

<!-- grade_id: cell-8f7be805bddd4762, points: 0, grade: false, solution: true -->
```python
data = [32, 42, 11, 15, 58, 44, 16]

result = io.StringIO()
with redirect_stdout(result):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-69e005cd46307fdb, points: 1, grade: true, solution: false -->
```python
# Testing cell
assert result.getvalue() == "[42, 11, 15, 58, 44, 16]\n"
```

**Q8** Print the length of the list `data`.

<!-- grade_id: cell-20ff0e0c86a745c9, points: 0, grade: false, solution: true -->
```python
data = [32, 42, 11, 15, 58, 44, 16]

result = io.StringIO()
with redirect_stdout(result):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a1a0a273991fb0a4, points: 1, grade: true, solution: false -->
```python
# Testing cell
assert result.getvalue() == "7\n"
```

**Q9** Print the third element of the list `data`.

<!-- grade_id: cell-2769c568b2695d86, points: 0, grade: false, solution: true -->
```python
data = [32, 42, 11, 15, 58, 44, 16]

result = io.StringIO()
with redirect_stdout(result):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0f56edb5b71e8344, points: 1, grade: true, solution: false -->
```python
# Testing cell
assert result.getvalue() == "11\n"
```

**Q10** Print the last element of the list `data`.

<!-- grade_id: cell-30bbcffba0b8be8d, points: 0, grade: false, solution: true -->
```python
data = [32, 42, 11, 15, 58, 44, 16]

result = io.StringIO()
with redirect_stdout(result):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-35bde65789e0106e, points: 1, grade: true, solution: false -->
```python
# Testing cell
assert result.getvalue() == "16\n"
```

**Q11** Print a list containing the last four elements of `data`.

<!-- grade_id: cell-ff997c6b4a2fc5eb, points: 0, grade: false, solution: true -->
```python
data = [32, 42, 11, 15, 58, 44, 16]

result = io.StringIO()
with redirect_stdout(result):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b59a34979e5d93d7, points: 1, grade: true, solution: false -->
```python
# Testing cell
assert result.getvalue() == "[15, 58, 44, 16]\n"
```

---

---
## Part C

How much time (in minutes/hours) did you spend on this homework assignment?

YOUR ANSWER HERE