# Denison CS181/DA210 Homework 2.c - Step 1

Before you turn this problem in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells** (in the menubar, select Kernel$\rightarrow$Restart And Run All).

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import pandas as pd
```

---

## Part A: Row Selection

**Q1:** Convert the following dictionary (a DoL representation) to a data frame called `testdf`.

```
    { 'Name'  : ['Ross', 'Rachel', 'Chandler'], 
      'Test1' : [82, 85, 95], 
      'Test2' : [89, 85, 90], 
      'Test3' : [93, 85, 87] }
```

<!-- grade_id: cell-3ec9b06343f7d5d6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# Display the DataFrame
testdf
```

<!-- grade_id: cell-bacb0663343d7f7c, points: 1, grade: true, solution: false -->
```python
# Testing cell
assert testdf.shape == (3,4)
```

**Q2:** Assign to variable `test1` the column vector (a `Series`) associated with the scores of **Test1**.

```python
# YOUR CODE HERE
raise NotImplementedError()

# Display the Series
test1
```

```python
# Testing cell
assert df.shape == (3,4)
assert isinstance(test1, pd.core.series.Series)
assert list(test1) == [82, 85, 95]
```

**Q3:** Assign to a variable `ross_rachel_df` the two-row `DataFrame` of Ross and Rachel's test scores.  You should select these rows using the numerical index.

<!-- grade_id: cell-15d0798550e374e5, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# Display the DataFrame
ross_rachel_df
```

```python
# Testing cell
assert ross_rachel_df.shape == (2,4)
assert isinstance(ross_rachel_df, pd.core.frame.DataFrame)
assert list(ross_rachel_df.Test1) == [82,85]
assert list(ross_rachel_df.Test2) == [89,85]
assert list(ross_rachel_df.Test3) == [93,85]
```

**Q4:** Assign to variable `rachel_df` the one-row `DataFrame` of Rachel's test scores.  You should filter using the `'Name'` column.

```python
# YOUR CODE HERE
raise NotImplementedError()

# Display the DataFrame
rachel_df
```

```python
# Testing cell
assert rachel_df.shape == (1,4)
assert isinstance(rachel_df, pd.core.frame.DataFrame)
```

**Q5:** Assign to variable `chandler_test3` the scalar integer of Chandler's third test score.

```python
# YOUR CODE HERE
raise NotImplementedError()

# Display the value
chandler_test3
```

```python
# Testing cell
assert chandler_test3.item() == 87
```

**Q6**: Compute a column Series, assigned to `final`, that equals the average of the three test scores.

<!-- grade_id: cell-f52532d3db1fb99c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# Display the Series
final
```

<!-- grade_id: cell-402ed44a99667b38, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert isinstance(final, pd.core.series.Series)
assert list(final)[:2] == [88.0, 85.0]
```

---

## Part B: Combining Projection and Selection

**Q7** Use the following data to create a DataFrame named `df`:
```
    data = {'animal':   ['cat','cat','snake','dog',
                         'dog','cat','snake','cat','dog','dog'],
            'age':      [2.5, 3, 0.5, 7, 5, 2, 4.5, 4, 7, 3],
            'visits':   [1, 3, 2, 3, 2, 3, 1, 1, 2, 1],
            'priority': ['yes','yes','no','yes','no', 
                         'no','no','yes','no', 'no']}
```

Using as index:
```
    labels = ['a', 'b', 'c', 'd', 'e', 
              'f', 'g', 'h', 'i', 'j']
```

Then, carry out a selection of rows `d` through `g`, inclusive, using row labels (`Index` elements), not positions.  Assign the result to `df2`.

<!-- grade_id: cell-61eb39ee4ca8447a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# Display the DataFrame
df2
```

<!-- grade_id: cell-f2ba02aa5279c53e, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert df.shape == (10,4)
assert df2.shape == (4,4)
assert 'd' in df2.index
assert 'g' in df2.index
```

**Q8:** In a single assignment statement, select the data in rows of `df` labeled `c`, `d`, and `e` and project the columns `'animal'` and `'age'`, assigning to `df3`.

<!-- grade_id: cell-37270ca1f13d69c2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# Display the DataFrame
df3
```

<!-- grade_id: cell-638458b2a637a019, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert df3.shape == (3, 2)
assert 'c' in df3.index
assert 'e' in df3.index
assert 'animal' in df3.columns
assert 'age' in df3.columns
```

**Q9:** Create a column vector `ageVisit` (a `Series`) which, for each row in `df`, is the age of the animal divided by the number of visits for the animal.

```python
# YOUR CODE HERE
raise NotImplementedError()

# Display the Series
ageVisit
```

```python
# Testing cell
assert isinstance(ageVisit, pd.core.series.Series)

assert len(ageVisit) == df.shape[0]
assert list(ageVisit[:3]) == [2.5, 1.0, 0.25]
```

**Q10:** Create a `DataFrame` `df4` with only the rows of `df` in which the `age` is greater than 3.

```python
# YOUR CODE HERE
raise NotImplementedError()

# Display the DataFrame
df4
```

```python
# Testing cell
assert isinstance(df4, pd.core.frame.DataFrame)

assert df4.shape == (5, 4)
assert list(df4.visits) == [3, 2, 1, 1, 2]

assert list(df4.index) == ['d', 'e', 'g', 'h', 'i']
```

**Q11:** From the original data frame, project the animal and visits columns from those rows where `priority` is `yes`.  Assign the result to `df5`.

```python
# YOUR CODE HERE
raise NotImplementedError()

# Display the DataFrame
df5
```

```python
# Testing cell
assert df5.shape == (4, 2)
assert "animal" in df5
assert "visits" in df5
assert "priority" not in df5

assert list(df5.index) == ['a', 'b', 'd', 'h']
```

---

---
## Part C

How much time (in minutes/hours) did you spend on this homework assignment?

YOUR ANSWER HERE