# Denison CS181/DA210 Homework 3.d - Step 1

Before you turn this problem in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells** (in the menubar, select Kernel$\rightarrow$Restart And Run All).

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

> **In the questions that follow, we are looking for XPath declarative solutions to the problems, not procedural solutions.  You will not get credit for procedural solutions.**

---

## Part A: Parse and familiarize yourself with the file

**Q1:** Begin by importing whatever modules you need, reading in and parsing the dataset in the file `bookstore.xml` in the `publicdata` directory.

You should name the variable representing the root node `books_root`.

```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8e1e34e520cb833f, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert type(books_root) is etree._Element
assert len(books_root) == 12
```

---

## Part B: Using XPath

**Q2:** Using `bookstore.xml`, create a Python list called `books` containing the **titles** of all books. Your list `books` should be a list of strings.

<!-- grade_id: cell-1d693d29dd846dd8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# Print the book titles
print(books)
type(books[0])
```

<!-- grade_id: cell-a8b4fd3cbd5c6d8d, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert len(books) > 0 and type(books[0]) is etree._ElementUnicodeResult
assert len(books)==12
assert 'Lover Birds' in books and 'Splish Splash' in books
```

**Q3:** Create a list, named `less`, of book ids that cost less than `$6`. Note that `id` is an attribute.

You can get the value of an attribute in XPath using `@attName`, where `attName` is the name of the attribute.

<!-- grade_id: cell-6e3c08a87e87befc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# Print the resulting list
less
```

<!-- grade_id: cell-caa0de2ac1828f14, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert len(less) > 0 and type(less[0]) is etree._ElementUnicodeResult
assert 'bk104' in less
assert 'bk101' not in less
assert len(less)==7
```

**Q4:** Create a list, `eva`, of book titles for which Eva Corets was the author. Your list `eva` should be a list of strings.

<!-- grade_id: cell-7c7d0674f1ec9421, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# Print the list
eva
```

<!-- grade_id: cell-331030ba75384a5c, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert len(eva) > 0 and type(eva[0]) is etree._ElementUnicodeResult
assert len(eva)==3
assert 'Maeve Ascendant' in eva
assert 'Paradox Lost' not in eva
```

**Q5:** Find the average book price for all books that are not fantasy in this file, assigning the resulting value to a variable `avgprice`.

**Hints** First, use XPath to get a list of the price strings (text) based on a single XPath query.  Then use a list comprehension to build a list of `float` values converting the strings to real-valued numbers.  Finally, perform the average based on the values and length of the list.

<!-- grade_id: cell-6d115874424475c8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# Print the average price
avgprice
```

<!-- grade_id: cell-5272bfd99dcd3cab, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert(avgprice > 23.82)
assert(avgprice < 24)
```

**Q6:** Create a list called `lessFantasy` containing the titles of the books where the price is under `$40` and not in the fantasy genre.

<!-- grade_id: cell-f712efddd8ca1dbd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# Print the list
lessFantasy
```

<!-- grade_id: cell-d20012db748daea5, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert len(lessFantasy)==6
assert 'Paradox Lost' in lessFantasy
assert 'Maeve Ascendant' not in lessFantasy
```

---

---

## Part C

**Q7:** How much time (in minutes/hours) did you spend on this homework assignment?

YOUR ANSWER HERE

**Q8:** Who was your partner for this assignment?  If you worked alone, say so instead.

YOUR ANSWER HERE